# Cours ENPC - Pratique du calcul scientifique

#### Instructions :

- Avant de rendre ce travail, assurez-vous que tout s'exécute correctement. Pour vérifier, **redémarrez VS Code** puis **exécutez toutes les cellules**.

- <mark style="color: red;">
    La plupart des exercices de ce notebook seront corrigés automatiquement.
    Vous pouvez modifier le contenu des cellules existantes et en ajouter de nouvelles.
    Cependant, pour garantir le bon déroulement de la correction automatique,
    <strong>ne supprimez ni ne modifiez les métadonnées des cellules fournies,</strong>
    en particulier celles qui nécessitent votre réponse. </mark>

- Assurez-vous de remplir tous les endroits indiqués par `YOUR CODE HERE` ou `YOUR ANSWER HERE`, ainsi que votre nom ci-dessous :

- Supprimez les lignes `error("No code provided")` une fois que vous avez complété votre solution.

```julia
NAME = "Prénom Nom"

# En soumettant ce notebook, vous acceptez la phrase suivante
println("Je certifie que ce notebook est le résultat de mon travail personnel. Signé : $NAME")
```

## Semaine 1 : Introduction à Julia

```julia
using Plots
using LaTeXStrings
```

### <font color='green'>Comprendre la sortie des cellules Jupyter</font>

- La sortie d'une série de commandes Julia est la dernière expression :

```julia
# Jupyter will display the result of `a + b` after running the cell

a = 5
b = 4
c = a + b
```

- L'utilisation de `;` en fin de commande supprime l'affichage

```julia
# Jupyter will not display anything

a = 5
b = 4
a + b;
```

- Cela s'applique aussi aux graphiques. Comparer les deux situations suivantes :

```julia
plot(cos)
```

```julia
plot(sin);
```

- On peut aussi utiliser `@show` ou `print` :

```julia
a = 5
b = 4
@show a
println(b)
a + b
```

### <font color='green'>Manipulation de chaînes de caractères</font>

- Les chaînes de caractères sont définies en Julia entre guillemets doubles `"`. Contrairement à Python, les guillemets simples `'` ne sont utilisés que pour les caractères simples.

```julia
s = "Hello world!"
```

```julia
'o' ∈ s # ∈ is an alias for in
```

- Contrairement à de nombreux langages qui utilisent `+`, Julia utilise `*` pour la concaténation de chaînes. L'idée est de souligner la non-commutativité de cette opération

```julia
s1 = "Hello " ; s2 = "world!"
s = s1 * s2
println(s)
```

```julia
print(s1^5)
```

- L'interpolation dans une chaîne s'obtient avec `$`

```julia
col = "red" ; x = 20_000 ; discount = 5
s = "My $col car cost $(x*(1-discount/100)) €"
print(s)
```

```julia
s = "Hello world!"
ordinal_suffix(i) = i == 1 ? "st" : i == 2 ? "nd" : "th"
for (i, c) ∈ enumerate(s) 
    println("The character at $i$(ordinal_suffix(i)) place is $c") 
end
```

> Notez que les chaînes LaTeX peuvent être construites grâce au package [LaTeXStrings](https://github.com/JuliaStrings/LaTeXStrings.jl), importé ci-dessus et utilisé plus loin dans ce notebook.

### <font color='green'>Utilisation et définition de fonctions</font>

#### <font color='orange'>[Exercice 1]</font>
Écrire une fonction qui retourne la somme des `n` premiers entiers au carré en utilisant une boucle `for`.

> *Indices :*
>  - lire la documentation de `function` (ouvrir un REPL et taper `?` suivi du mot recherché, ici `function`),
>  - un intervalle d'entiers de `m` à `n` s'écrit simplement `m:n` <mark>(contrairement à Python, les bornes sont incluses dans l'intervalle)</mark>.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```julia
function sum_of_squares(n)
    # YOUR CODE HERE [⚠ DO NOT DELETE THIS CELL]
    error("No code provided")
end;
```

<!-- grade_id: squares-tests, points: 2, grade: true, solution: false -->
```julia
@assert sum_of_squares(1) == 1 "❌"
@assert sum_of_squares(2) == 5 "❌"
@assert sum_of_squares(5) == 55 "❌"
@assert sum_of_squares(10) == 385 "❌"
```

**Remarques :**
- La même fonction peut s'écrire sous forme compacte

```julia
short_sum_of_squares(n) = (1:n)'*(1:n)

@show short_sum_of_squares(2)
@show short_sum_of_squares(4);
```

- Lire la documentation de la fonction `sum` pour comprendre pourquoi le code suivant fonctionne aussi.

```julia
other_sum_of_squares(n) = sum(x -> x^2, 1:n)

@show other_sum_of_squares(2)
@show other_sum_of_squares(4);
```

- La fonction peut être appliquée élément par élément à un tableau en utilisant l'opérateur de diffusion `.`

```julia
@show short_sum_of_squares.([2, 4, 5])

# The above command is syntactic sugar for
@show broadcast(short_sum_of_squares, [2, 4, 5]);
```

- La diffusion fonctionne même pour des fonctions à plusieurs arguments

```julia
geometric_mean(a, b) = √(a*b)
@show geometric_mean.(2, [8, 18]);
```

- Il est parfois commode de définir des fonctions *anonymes* avec la notation `->` :

```julia
# Plot the sum of cubes for different n
scatter(n -> sum(x -> x^3, 1:n), 1:5, label="sum of cubes")
```

#### <font color='orange'>[Exercice 2]</font>

La méthode babylonienne pour calculer la racine carrée de $a > 0$ est donnée par l'itération
$$
\forall k\in\mathbb{N},\qquad x_{k+1}=\frac{1}{2}\left(x_k+\frac{a}{x_k}\right)
$$

Écrire une fonction `my_sqrt(a, x₀)` qui retourne une approximation de la racine carrée de `a` par 10 itérations de cette méthode, en partant de `x₀`.

<!-- grade_id: baby, points: 0, grade: false, solution: true -->
```julia
function my_sqrt(a, x₀)
    # YOUR CODE HERE [⚠ DO NOT DELETE THIS CELL]
    error("No code provided")
end;
```

<!-- grade_id: baby-tests, points: 2, grade: true, solution: false -->
```julia
@assert my_sqrt(4, 1) ≈ 2 "❌"
@assert my_sqrt(144, 1) ≈ 12 "❌"
```

### <font color='green'>Tracé de courbes et création d'animations</font>

- La fonction `plot` peut être appelée de différentes façons.
  Une première approche consiste à lui passer une fonction :

```julia
f(x) = cos(x) * exp(-abs(x))
plot(f, xlims=(-5, 5))
```

- Une autre option, souvent plus flexible, est de lui passer des tableaux $x$ et $y$ :

```julia
f(x) = cos(x) * exp(-abs(x))
x = -5:.01:5
plot(x, f.(x))
```

- Utiliser les commandes suivies de `!` pour modifier un graphique existant.

```julia
plot(cos, label = "f1")
plot!(sin, label = "f2")
xlims!(-5, 5)
title!("title")
xlabel!("x")
ylabel!("y")
```

- Les options peuvent être regroupées dans un appel à `plot!` :

```julia
plot(cos, label = "f1")
plot!(sin, label = "f2")
plot!(xlims = (-5, 5), title = "title", xlabel = "x", ylabel = "y")
```

- Les animations peuvent être créées de la façon suivante :

```julia
anim = Animation()
for θ in LinRange(0, 2π, 50)
    plot(x -> cos(x - θ), label=nothing)
    title!("θ = $(round(θ, digits=3))")
    frame(anim)
end
gif(anim, fps=10)
```

- La bibliothèque `Plots` définit la macro `@animate` pour simplifier la création d'animations :

```julia
anim = @animate for θ in LinRange(0, 2π, 50)
    plot(x -> cos(x - θ), label=nothing)
    title!("θ = $(round(θ, digits=3))")
end
gif(anim, fps=10)
```

#### <font color='orange'>[Exercice 3]</font>
Créer une animation illustrant le mouvement d'une balle de golf dont la position dans le plan $x,y$ est donnée par
$$
\begin{aligned}
x(t) &= u_0 t, \\
y(t) &= v_0 t - g \frac{t^2}{2},
\end{aligned}
$$
Utiliser les paramètres fournis ainsi que la fonction `scatter`,
et rendre l'animation aussi soignée que possible.

<!-- grade_id: golf, points: 4, grade: true, solution: true -->
```julia
u₀ = 50
v₀ = 50
g = 9.81
# YOUR CODE HERE [⚠ DO NOT DELETE THIS CELL]
error("No code provided")
```

### <font color='green'>Formats à virgule flottante</font>

Seul un nombre fini de réels peut être représenté exactement sur un ordinateur.
Par conséquent, certains réels sont représentables exactement et d'autres non.
La norme IEEE 754, à laquelle se conforment pratiquement tous les langages de programmation,
spécifie un petit nombre de **formats à virgule flottante**.
L'ensemble des réels représentables dans chaque format est de la forme
$$
\begin{align*}
    F(p, E_{\min}, E_{\max})
    = \Bigl\{ & (-1)^s 2^E (b_0. b_1 b_2 \dots b_{p-1})_2 \colon \\
              & \qquad s \in \{0, 1\}, b_i \in \{0, 1\} \, \text{et} \, E_{\min} \leq E \leq E_{\max} \Bigr\}.
\end{align*}
$$
Ici $(b_0.b_1 b_2 \dots b_{p-1})_2$ est un nombre en représentation binaire, c'est-à-dire
$$
(b_0.b_1 b_2 \dots b_{p-1})_2 = b_0 + \frac{1}{2} b_1 + \dotsc + \frac{1}{2^{p-1}} b_{p-1}.
$$
Pour un élément $x \in F(p, E_{\min}, E_{\max})$, $(b_0. b_1 b_2 \dots b_{p-1})_2$ est appelée la *mantisse*, $s$ est le *signe* et $E$ est l'*exposant*.
Trois paramètres apparaissent dans la définition de l'ensemble ci-dessus :
- $p$ est le nombre de bits significatifs (appelé aussi la **précision**),
- $E_{\min}$ est l'exposant minimum,
- $E_{\max}$ est l'exposant maximum.

À partir de la précision, l'**epsilon machine** associé au format est défini comme
$$
\varepsilon = 2^{-(p-1)}
$$

Les paramètres des formats les plus couramment utilisés spécifiés par la norme sont les suivants :

| Paramètre | Demi-précision (`Float16`) | Simple précision (`Float32`) | Double précision (`Float64`) |
| --- | --- | --- | --- |
| $p$ | 11 | 24 | 53 |
| $E_{\min}$ | -14 | -126 | -1022 |
| $E_{\max}$ | 15 | 127 | 1023 |

En Julia, les formats à virgule flottante demi, simple et double précision sont donnés par
`Float16`, `Float32` et `Float64`, le format par défaut étant `Float64`.

```julia
a = √2
b = Float32(√2)
c = Float16(√2)

@show a
@show b
@show c

@show typeof(a)
@show typeof(b)
@show typeof(c);
```

L'epsilon machine doit être compris comme une mesure de l'espacement relatif entre un nombre à virgule flottante et le suivant, au sein d'un format.
En effet, le nombre 1 est représentable exactement dans tous les formats,
et le prochain nombre représentable est précisément $1 + \varepsilon$

```julia
a = Float16(1.0)
@show nextfloat(a) - a
@show eps(Float16);
```

De même, le prochain nombre représentable après $2$ est $2 + 2\varepsilon$
et le prochain nombre représentable après $4$ est $4 + 4\varepsilon$ :

```julia
ε = eps(Float64)
@show nextfloat(2.0) - 2.0
@show 2ε;
@show nextfloat(4.0) - 4.0
@show 4ε;
```

Comme le suggèrent leurs noms,
un nombre `Float16` est stocké sur 16 bits,
un `Float32` sur 32 bits,
et un nombre `Float64` sur 64 bits.
Parmi ces bits, un est utilisé pour stocker le signe,
et les autres pour stocker la mantisse et l'exposant.
De plus, quelques combinaisons de bits sont réservées pour représenter les valeurs spéciales
`Inf`, `-Inf` et `NaN`, ce dernier étant l'acronyme de « not a number » (pas un nombre).

Un **octet**, abrégé par un B majuscule, est une unité composée de 8 bits.
C'est l'unité généralement utilisée pour décrire la quantité de mémoire disponible dans les produits commerciaux ;
par exemple, un ordinateur portable dispose typiquement de 8 Go ou 16 Go de mémoire vive.
Il est souvent utile d'avoir une idée approximative de l'espace occupé en mémoire par les objets mathématiques.
Par exemple, une matrice 1000 × 1000 de nombres `Float64` nécessite $8 \times 10^6$ octets pour être stockée,
et une matrice 10000 × 10000 de nombres `Float64` requiert près de 800 Mo de mémoire vive.

```julia
A = zeros(10_000, 10_000)
@show Base.summarysize(A);
```

Pour approfondir notre compréhension des formats à virgule flottante,
traçons sur un même graphique l'espacement entre
les nombres successifs `Float16`, `Float32` et `Float64`
dans l'intervalle $[1, 10^4]$. On constate que la densité de
nombres à virgule flottante représentables autour de $x$ diminue
avec $x$.

```julia
n = 1000
x_16 = (^).(10, LinRange{Float16}(0, 4, n))
x_32 = (^).(10, LinRange{Float32}(0, 4, n))
x_64 = (^).(10, LinRange{Float64}(0, 4, n))
spacing16 = nextfloat.(x_16) - x_16
spacing32 = nextfloat.(x_32) - x_32
spacing64 = nextfloat.(x_64) - x_64
plot(x_16, spacing16, label="Float16", )
plot!(x_32, spacing32, label="Float32")
plot!(x_64, spacing64, label="Float64")
plot!(xscale=:log10, yscale=:log10,
      xlabel="x", ylabel="Δx", legend=:bottomright)
```

#### <font color='orange'>[Exercice 4]</font>
Pour approfondir votre compréhension,
- Lire la documentation de la fonction `nextfloat`.
- En utilisant cette fonction,
  tracer l'espacement **relatif**, normalisé par l'epsilon machine du format `Float32`,
  entre les nombres `Float32` successifs dans l'intervalle $[10^{-2}, 10^2]$.
  Plus précisément, tracer la fonction
  $$
  x \mapsto \frac{\text{nextfloat}(x) - x}{x \varepsilon}
  $$

Utiliser une échelle logarithmique pour l'axe $x$ uniquement.
Il peut être utile d'utiliser `LinRange{Type}(a, b, n)`
pour créer un vecteur de $n$ nombres de type `Type` équidistants entre $a$ et $b$.

<!-- grade_id: spacing, points: 4, grade: true, solution: true -->
```julia
# YOUR CODE HERE [⚠ DO NOT DELETE THIS CELL]
error("No code provided")
```

Abordons maintenant une question naturelle découlant de cette discussion :
que se passe-t-il si le résultat d'une opération mathématique n'est pas exactement représentable dans le format à virgule flottante ?
La réponse est simple : le résultat est arrondi au nombre représentable le plus proche.
Cela explique les résultats surprenants suivants :

```julia
@show 1 + eps()/3      # the result is 1
@show 1 + 2*eps()/3    # the result is 1 + ε
@show .1 + .2 == .3    # this is False
```

Les erreurs causées par ce processus d'arrondi sont appelées **erreurs d'arrondi**.
Elles sont particulièrement importantes pour le calcul des dérivées.
En effet, supposons que l'on souhaite calculer,
pour une fonction différentiable $f\colon \mathbb R \to \mathbb R$, une approximation de $f'(1)$.
L'approche la plus naturelle est de calculer pour un petit $\delta > 0$
$$
d(\delta) = \frac{f(1+ \delta) - f(1)}{\delta}.
$$
En arithmétique exacte, cela fournirait une approximation de $f'(1)$ avec une erreur en $\mathcal O(\delta)$,
en supposant $f''(1) \neq 0$.
En arithmétique d'ordinateur, cependant, une erreur supplémentaire apparaît en raison des arrondis,
dont on peut montrer qu'elle est de l'ordre de $\mathcal O(\varepsilon/\delta)$.
Le meilleur choix est de prendre $\delta \approx \sqrt{\varepsilon}$,
de sorte que les erreurs soient approximativement du même ordre de grandeur.

```julia
f(x) = exp(x)
d(δ) = (f(1+δ) - f(1))/δ
δs = 10 .^(0:-.1:-17)
err = abs.(d.(δs) .- exp(1))
plot(δs, err, xscale=:log10, yscale=:log10, label=L"|d(\delta) - f'(1)|", size=(900,450))
plot!([eps()], seriestype=:vline, label = L"\varepsilon")
plot!([sqrt(eps())], seriestype=:vline, label = L"\sqrt{\varepsilon}")
```

Pour des $\delta$ très petits de l'ordre de l'epsilon machine $\varepsilon$,
l'approximation obtenue est complètement fausse.
Comprenez-vous les observations suivantes ?

```julia
δ = eps()/4; @show (exp(δ) - exp(0)) / δ
δ = eps()/2; @show (exp(δ) - exp(0)) / δ
δ = eps(); @show (exp(δ) - exp(0)) / δ;
```

**Remarque**. Dans de nombreux contextes, il est utile d'utiliser des nombres à virgule flottante avec plus de précision que les types standard,
afin de réduire les erreurs d'arrondi et d'observer l'erreur intrinsèque,
liée au caractère inexact de nombreux algorithmes même en arithmétique exacte.
Le type `BigFloat` en Julia est un format à précision arbitraire.
Le nombre de bits significatifs correspondant peut être défini à l'aide de la fonction `setprecision`.

```julia
@show Float64(π)    

setprecision(100, base=10)  # We want 100 digits in base 10
BigFloat(π)
```

### <font color='green'>Manipulation des tableaux</font>

- Définition explicite simple

```julia
u = [1, 2, 3]  # defines a column Vector
display(u)
@show typeof(u)
```

- Remarquer la différence avec

```julia
u = [1 2 3] # defines a Matrix of 1 line
display(u)
@show typeof(u)
```

- Définition explicite d'une matrice

```julia
A = [1 2 3
     4 5 6
     7 8 9]
```

```julia
A = [1 2 3; 4 5 6; 7 8 9]
```

- Extraction d'une colonne ou d'une ligne
> **Note :** en `Julia`, les indices commencent à `1` comme en `Fortran` (**et non à `0` comme en `Python` ou `C`**)

```julia
A[1,:]
```

```julia
A[:,1]
```

- `:` est un raccourci pour `begin:end` ou `1:end`

```julia
A[2,1:end]
```

```julia
A[begin:end,2]
```

- Construction d'un tableau par syntaxe de compréhension

```julia
A = [10i+j for i in 0:2, j in 0:3]
display(A)
print("\nA = $A\n\n")
@show A
@show A[1,1]
@show A[1,2]
@show length(A)
@show size(A)
@show size(A, 1)
@show size(A, 2)
@show eltype(A)
;
```

- Matrice adjointe

```julia
A = [1 2 3
     4 5 6]
A'
```

```julia
display(A'A)
display(A*A')
```

```julia
A = [2+3im 4+2im 1+3im
     1+2im 2+3im 5+6im]
A'
```

```julia
display(A'A)
display(A*A')
```

- Concaténation de tableaux

```julia
A = [1 2 3
     4 5 6]
u = [10, 20, 30]
v = [100, 200]
;
```

```julia
[A v]
```

```julia
[A; u']
```

```julia
[A A]
```

```julia
[A; A]
```

```julia
hcat(A, A, A)
```

```julia
vcat(A, A, A)
```

- Les composantes sont stockées par colonnes et peuvent être accédées par un indice simple même pour des tableaux multidimensionnels
- `eachindex` et `CartesianIndices` permettent d'itérer sur tous les indices d'un tableau (indice simple pour le premier, multi-indice pour le second)
- Un `CartesianIndex` peut être converti en `Tuple`
- `...` est l'opérateur splat (voir la documentation)

```julia
A = [1 2 3; 4 5 6]

for k ∈ eachindex(A)
    println("A[$k] = ", A[k])
end
println()

for k ∈ CartesianIndices(A)
    println("A[$k] = A$(Tuple(k)...) = ", A[k])
end
println()

for i in 1:size(A, 1)
    for j in 1:size(A, 2)
        println("A[$i, $j] = ", A[i, j])
    end
end
println()

for i in 1:size(A, 1), j in 1:size(A, 2)
    println("A[$i, $j] = ", A[i, j])
end
```

- Une matrice (tableau d'ordre 2) n'est pas un vecteur de vecteurs

```julia
B = [[10i+j for i in 0:2] for j in 0:3]
display(B)
@show B isa Matrix 
;
```

- Les tableaux sont mutables (si le type est cohérent), les tuples ne le sont pas !

```julia
v = [1, 2]
@show typeof(v)
@show eltype(v)
@show v
v[1] = 3
@show v
# v[1] = 1.5 # error since v is of type Vector{Int64}
# @show v
t = (1, 2)
@show t[1] ;
# t[1] = 3 # error since t is not mutable
```

- Définition d'un tableau vide et construction *a posteriori*

> Remarquer la différence entre `push!` et `append!` (il est rappelé que pour les deux, `!` exprime que par convention, et non par obligation, l'argument est modifié)

```julia
v = [] # void vector containing any type (not optimal!)
@show typeof(v)
push!(v, 1)
append!(v, [2.5, 3.6]) # pushes each item of the vector [2.5, 3.6]
@show v
push!(v, [4.1, 5.2]) # pushes a vector as new item
@show v ;
@show v[end]
@show v[end-1] ;
```

- Pour de meilleures performances, il est préférable, si possible, de préciser le type des éléments du vecteur lors de la construction

```julia
v = Float64[]
@show typeof(v)
push!(v, 1)
append!(v, [2//3, π]) # pushes each item of the vector [2//3, π] and convert on-the-fly to Float64
@show v
# push!(v, [4.3, 5.7]) # throws an error since [4.3, 5.7] is not of type Float64
;
```

- Référence d'un tableau

```julia
A = [10i+j for i in 1:2, j in 1:3]
@show A
B = A # B is a reference to A (both point to the same memory location)
@show B == A # same content
@show B === A # same memory location
B .= 99 # .= assigns all the coefficients of the array to 99
@show B
@show A # A is of course also affected
;
```

- Copie d'un tableau

```julia
A = [10i+j for i in 1:2, j in 1:3]
@show A
B = copy(A) # B and A are two distinct arrays
@show B == A # same content
@show B === A # not the same memory location
B .= 99 # .= assigns all the coefficients of the array to 99
@show B
@show A # A unaffected
;
```

- Tranche en tant que l-value ≡ **référence**

```julia
A = [10i+j for i in 1:6, j in 1:8]
display(A)
A[1:2:end, 2:2:end] .= 0 # components of odd line and even column are set to 0
display(A)
;
```

- Tranche en tant que r-value ≡ **copie**

```julia
A = [10i+j for i in 1:6, j in 1:8]
display(A)
B = A[1:2:end, 2:2:end] # B is a copy of the submatrix
display(B)
B .= 0
display(B)
display(A) # A remains unaffected
;
```

- Référence en tant que r-value via `view`

```julia
A = [10i+j for i in 1:6, j in 1:8]
display(A)
B = view(A, 1:2:size(A,1), 2:2:size(A,2)) # B is a view of the submatrix (reference not a copy)
display(B)
B .= 0
display(B)
display(A) # A is modified
;
```

- Autres façons de définir des tranches par sélection d'indices

```julia
A = [10i+j for i in 1:6, j in 1:8]
A[isodd.(1:end), iseven.(1:end)] .= 0 ; display(A)
A = [10i+j for i in 1:6, j in 1:8]
A[findall(isodd,A)] .= 0 ; display(A) ;
```

- Expliquer les lignes suivantes

```julia
A = [10i+j for i in 1:6, j in 1:8]
B = view(A, collect(1:size(A,1)) .% 3 .== 0, collect(1:size(A,2)) .% 5 .!= 0)
B .= 0
A
```

```julia
A = rand(5, 5) ; display(A)
A[findall(x -> x < 1/2, A)] .= 0
A
```

```julia
n, m = 7, 9
B = rand(n, m)
u = rand(m)
display(B*u)
subi = collect(1:n) .% 3 .!= 0 ; @show subi
subj = collect(1:m) .% 3 .== 0 ; @show subj
B[subi,subj]*u[subj]
```

#### <font color='orange'>[Exercice 5]</font>

Écrire une fonction qui calcule la comatrice d'une matrice explicitement à partir du déterminant de sous-matrices.

On rappelle que la composante $C_{ij}$ de la comatrice de $A$ est obtenue comme le déterminant de la sous-matrice obtenue à partir de $A$ en supprimant sa $i^\textrm{ème}$ ligne et sa $j^\textrm{ème}$ colonne, multiplié par $(-1)^{i+j}.$ Le déterminant d'une petite matrice dense est calculé par la fonction `det` disponible dans la bibliothèque `LinearAlgebra`.

<!-- grade_id: comatrix, points: 0, grade: false, solution: true -->
```julia
using LinearAlgebra
function comatrix(A)
    C = zero(A)
    # YOUR CODE HERE [⚠ DO NOT DELETE THIS CELL]
    error("No code provided")
    return C
end
```

<!-- grade_id: comatrix-test, points: 4, grade: true, solution: false -->
```julia
n = 10
A = rand(n, n)
C = comatrix(A)
display(C'A)
@assert C'A ≈ det(A)*LinearAlgebra.I
```

### <font color='green'>Création de structures personnalisées</font>

Julia n'est pas un langage orienté objet.
Il supporte à la place la *distribution multiple* (multiple dispatch),
dont la motivation est exposée éloquemment dans la [documentation de Julia](https://docs.julialang.org/en/v1/manual/methods/) :

> The choice of which method to execute when a function is applied is called dispatch. Julia allows the dispatch process to choose which of a function's methods to call based on the number of arguments given, and on the types of all of the function's arguments. This is different from traditional object-oriented languages, where dispatch occurs based only on the first argument, which often has a special argument syntax, and is sometimes implied rather than explicitly written as an argument. Using all of a function's arguments to choose which method should be invoked, rather than just the first, is known as multiple dispatch. Multiple dispatch is particularly useful for mathematical code, where it makes little sense to artificially deem the operations to "belong" to one argument more than any of the others: does the addition operation in x + y belong to x any more than it does to y? The implementation of a mathematical operator generally depends on the types of all of its arguments. Even beyond mathematical operations, however, multiple dispatch ends up being a powerful and convenient paradigm for structuring and organizing programs.

```julia
struct Dog
  name::String
end

struct Cat
  name::String
end

greet(a::Dog, b::Dog) = "$(a.name): Wouf wouf $(b.name), content de te revoir mon vieux"
greet(a::Dog, b::Cat) = "$(a.name): Wouf wouf wouf wouf !"
greet(a::Cat, b::Dog) = "$(a.name) s'éloigne discrètement de $(b.name)"
greet(a::Cat, b::Cat) = "$(a.name) commence à ronronner"

milou = Dog("Milou")
rantanplan = Dog("Rantanplan")
miaouss = Cat("Miaouss")
mrMoustache = Cat("MrMoustache")

@show greet(milou, rantanplan)
@show greet(milou, mrMoustache)
@show greet(mrMoustache, miaouss)
@show greet(mrMoustache, rantanplan);
```

Pour conclure cette séance,
illustrons la puissance de la distribution multiple avec deux autres exemples.
Premièrement, créons un type spécifique pour la matrice identité,
et redéfinissons les opérations `+` et `*` sur les matrices pour qu'elles fonctionnent avec ce nouveau type.

```julia
struct IdentityMatrix end

# This is required to redefine + and *
import Base: +,*

function +(A::Matrix, Id::IdentityMatrix)
    result = copy(A)
    for i in 1:size(A, 1)
        result[i, i] += 1
    end
    return result
end

+(Id::IdentityMatrix, A::Matrix) = +(A, Id)
*(A::Matrix, Id::IdentityMatrix) = copy(A)
*(Id::IdentityMatrix, A::Matrix) = copy(A)

A = ones(3, 3)
Id = IdentityMatrix()

# We can now write operations such as
@show A
@show A + Id
@show Id + A
@show A * Id;
```

Remarquer que nous n'avons pas eu besoin de modifier l'implémentation interne des matrices.
On notera également qu'avec notre implémentation de la matrice identité,
le produit `A*I` est pratiquement gratuit en temps de calcul.
Si la matrice identité avait été stockée comme une matrice pleine,
beaucoup de puissance de calcul aurait été gaspillée dans la multiplication matricielle.

```julia
n = 5000
A = rand(n, n)
I1 = IdentityMatrix()
I2 = [(i == j)*1.0 for i in 1:n, j in 1:n]

@time A*I1;
@time A*I2;
```

#### <font color='orange'>[Exercice 6]</font>
L'objectif de cet exercice est de créer une structure pour représenter les nombres duaux.
Les nombres duaux sont une construction mathématique qui étend
les réels en introduisant un nouvel élément, souvent
noté $\varepsilon$, tel que $\varepsilon^2 = 0$.
Un nombre dual est composé d'une partie réelle et d'une partie duale,
où la partie duale se comporte comme une quantité infinitésimale.
Ils seront très utiles plus loin dans le cours,
lorsque nous étudierons la différentiation automatique, une technique
pour calculer efficacement des dérivées numériquement.
Compléter la définition de structure ci-dessous en ajoutant les opérations `-`, `*` et `/` aux nombres duaux.

<!-- grade_id: dual, points: 0, grade: false, solution: true -->
```julia
import Base: +, -, *, /, show

struct Dual
    real::Float64
    dual::Float64
end

+(x::Dual, y::Dual) = Dual(x.real + y.real, x.dual + y.dual)
# YOUR CODE HERE [⚠ DO NOT DELETE THIS CELL]
error("No code provided")

show(io::IO, x::Dual) = print(io, x.real, x.dual ≥ 0 ? " + $(x.dual)" : " - $(-x.dual)", " ε")

# Example usage
a = Dual(3.0, 2.0)  # 3 + 2ε
b = Dual(1.0, 4.0)  # 1 + 4ε

@show a + b
@show a - b
@show a * b
@show a / b;
```

<!-- grade_id: dual-tests, points: 4, grade: true, solution: false -->
```julia
@assert a + b == Dual(4, 6)
@assert a - b == Dual(2, -2)
@assert a * b == Dual(3, 14)
@assert a / b == Dual(3, -10)
```

```julia

```